## Εργασία 2 ##

Καλωσήρθατε στην δεύτερη σας εργασία. Η εργασία αυτή έχει σκοπό να σας βοηθήσει να εμπεδώσετε τα δενδρικά μοντέλα και την αξιολόγηση μοντέλων.

Στην εργασία αυτή θα πρέπει να συμπληρώσετε κώδικα Python 3 στα σημεία που αναφέρουν # YOUR CODE HERE. Μην τροποποιείτε τον κώδικα που βρίσκεται εκτός αυτών των περιοχών.

Πρωτού παραδόσετε την εργασία σας σιγουρευτείτε ότι ο κώδικας σε όλα τα κελιά τρέχει σωστά. Για το σκοπό αυτό επιλέξτε από το μενού Χρόνος εκτέλεσης (runtime) -> Επανεκίνηση περιόδου λειτουργίας και εκτέλεση όλων.

Συμπληρώστε το όνομα (NAME) και το AEM σας παρακάτω:

```python
NAME = "Αλεξανδρίδης Φώτιος"
AEM = "9953"
```

**1** Διαβάστε το διαθέσιμο από την sklearn σύνολο δεδομένων diabetes που αφορά την πρόβλεψη της πορείας που θα έχει ένας ασθενής με διαβήτη ένα χρόνο μετά (εργασία παλινδρόμησης) και αποθηκεύστε το σε μια μεταβλητή με το όνομα diabetes. (0.5 μονάδες)

<!-- grade_id: cell-ac5d304e50a110a6, points: 0, grade: false, solution: true -->
```python
from sklearn.datasets import load_diabetes
diabetes = load_diabetes()
```

<!-- grade_id: cell-2d158133c851eb4e, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθής ανάγνωσης του συνόλου δεδομένων"""
assert diabetes.feature_names == ['age', 'sex', 'bmi', 'bp', 's1', 's2', 's3', 's4', 's5', 's6']
```

**2** Χωρίστε το σύνολο που διαβάσατε σε υποσύνολο ελέγχου 20% και υποσύνολο εκπαίδευσης 80% με την μέθοδο train_test_split και σπόρο τυχαιότητας 42. Aποθηκεύστε το σύνολο εκπαίδευσης σε μεταβλητές X_train, y_train και το σύνολο ελέγχου σε μεταβλητές X_test, y_test. (0.5 μονάδες)

<!-- grade_id: cell-a637aedb66ae6d89, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split
X_train, X_test, y_train, y_test = train_test_split(diabetes["data"],  diabetes["target"], test_size=0.2, random_state=42)
```

<!-- grade_id: cell-300331f01f617c3f, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == 0.02736
assert round(X_test[0][8], 5) == 0.03243
```

**3** Χρησιμοποιώντας τα δεδομένα εκπαίδευσης, εκπαιδεύστε δύο δένδρα παλινδρόμησης με κριτήρια διαχωρισμού τα squared_eror και poisson ([δείτε ενότητα 1.10.7.2](https://scikit-learn.org/stable/modules/tree.html#tree)) αντίστοιχα, θέτοντας τον σπόρο τυχαιότητας στην τιμή 42 και το μέγιστο βάθος στην τιμή 3. Υπολογίστε τη μετρική r2 με βάση τις προβλέψεις τους στα δεδομένα ελέγχου, τις οποίες αποθηκεύστε στις μεταβλητές y_pred1 (squared_error) και y_pred2 (poisson) αντίστοιχα.
Αποθηκεύστε τις μετρικές στις μεταβλητές r2_squared_error και r2_poisson αντίστοιχα. (1 μονάδα)

<!-- grade_id: cell-62cfb517c0f9271e, points: 0, grade: false, solution: true -->
```python
from sklearn.tree import DecisionTreeRegressor
from sklearn.metrics import r2_score
# Create the two models
regr_sq = DecisionTreeRegressor(max_depth=3, random_state=42, criterion="squared_error")
regr_ps = DecisionTreeRegressor(max_depth=3, random_state=42, criterion="poisson")
# Fit the models
regr_sq.fit(X_train, y_train)
regr_ps.fit(X_train, y_train)
# Calculate predictions
y_pred_sq = regr_sq.predict(X_test)
y_pred_ps = regr_ps.predict(X_test)
# Calculate r2 
r2_squared_error = r2_score(y_test, y_pred_sq)
r2_poisson = r2_score(y_test, y_pred_ps)
# Assign variables
y_pred1 = y_pred_sq
y_pred2 = y_pred_ps
```

<!-- grade_id: cell-714ce0b2b256307e, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθού υπολογισμού των μετρικών"""
assert round(r2_squared_error, 3) == 0.329
assert round(r2_poisson, 3) == 0.391
assert len(y_pred1) == len(y_pred2)
```

**4** Αναπτύξτε κώδικα με βάση την ιδιότητα [tree_](https://scikit-learn.org/stable/auto_examples/tree/plot_unveil_tree_structure.html#sphx-glr-auto-examples-tree-plot-unveil-tree-structure-py) του παρακάτω μοντέλου, προκειμένου να εμφανίσετε το μονοπάτι που ακολουθήθηκε στο δένδρο για την πρόβλεψη της πρώτης περίπτωσης του συνόλου ελέγχου. Σε περίπτωση που ελέγχεται δεύτερη φορά μια μεταβλητή στο μονοπάτι, θα πρέπει να διατηρούμε τον τελευταίο και στενότερο περιορισμό. Ακολουθούν παραδείγματα ορθής λειτουργίας, στα οποία φαίνεται και η μορφή αν-τότε που πρέπει να έχει το μονοπάτι. (4 μονάδες)

```python
from sklearn.tree import export_text
model = DecisionTreeRegressor(max_depth=3, random_state=42)
model.fit(X_train, y_train)
```

Output:
```
DecisionTreeRegressor(max_depth=3, random_state=42)
```

<!-- grade_id: cell-1697542afe8f99a2, points: 0, grade: false, solution: true -->
```python
def get_path(tree, instance, feature_names):
    # If passing one instance, add a dimension for correct shapes
    if len(instance.shape) == 1:
        instance = instance.reshape(1, -1)
    
    # Obtain threshold and feature values
    threshold = tree.tree_.threshold
    feature = tree.tree_.feature

    # Obtain the node indices for the decision path
    node_indicator = tree.decision_path(instance)
    leaf_id = tree.apply(instance)
    node_index = node_indicator.indices[
        node_indicator.indptr[0] : node_indicator.indptr[1]
    ]

    # Create a list to hold all rules
    rules = []

    # Iterate over all nodes
    for node_id in node_index:
        # continue to the next node if it is a leaf node
        if leaf_id[0] == node_id:
            continue

        # Obtain feature name
        feature_name = feature_names[feature[node_id]]
        # Obtain feature threshold
        feature_threshold = threshold[node_id]
        # Obtain feature value
        feature_value = instance[0, feature[node_id]]

        # check if value of the split feature for sample 0 is below threshold
        if feature_value <= feature_threshold:
            threshold_sign = "<="
        else:
            threshold_sign = ">"
        
        # Iterate over rules list, finding if this rule is before, further constricting it if needed
        found_rule = False
        for rule_idx in range(len(rules)):
            # Get the current rule
            rule = rules[rule_idx]
            # If a check is found
            if rule[0] == feature_name and rule[1] == threshold_sign:
                # To ensure addition does not happen later
                found_rule = True
                # Update the threshold, performing the correct check
                # This is redundand in the decision tree, but added for clarity
                if threshold_sign == "<=" and feature_threshold <= rule[2]:
                    rules[rule_idx][2] = feature_threshold
                if threshold_sign == ">" and feature_threshold > rule[2]:
                    rules[rule_idx][2] = feature_threshold
            
        if not found_rule:
            rules.append([feature_name, threshold_sign, feature_threshold])

    
    # Construct result text
    result_text = ""

    # Loop over all found rules
    for rule in rules:
        if result_text == "":
            starting_text = "αν "
        else:
            starting_text = "και "
        
        string_segment = starting_text + rule[0] + " " + rule[1] + " " + str(round(rule[2], 2)) + " "
        result_text += string_segment

    # Add final value
    result_text += "τότε " + str(round(tree.predict(instance)[0], 2))
    return(result_text)
```

<!-- grade_id: cell-8aabb13a60a743dd, points: 4, grade: true, solution: false -->
```python
"""Τεστ ορθού υπολογισμού του μονοπατιού"""
assert get_path(model, X_test[0], diabetes.feature_names) == 'αν bmi <= 0.01 και s5 > 0.01 και s4 <= 0.09 τότε 159.57'
assert get_path(model, X_test[1], diabetes.feature_names) == 'αν bmi > 0.01 και bmi <= 0.07 και s6 <= 0.03 τότε 175.8'
assert get_path(model, X_test[6], diabetes.feature_names) == 'αν bmi > 0.07 και s2 > 0.02 τότε 225.75'
```

**5** Διαχωρίστε τα δεδομένα diabetes σε υποσύνολα εκπαίδευσης 60% (X_train, y_train), επικύρωσης 15% (X_val, y_val) και ελέγχου 25% (X_test, y_test) με σπόρο τυχαιότητας 42. (1 μονάδα)

<!-- grade_id: cell-c2448c38ed5250e1, points: 0, grade: false, solution: true -->
```python
X_train, X_test, y_train, y_test = train_test_split(diabetes["data"],  diabetes["target"], test_size=0.25, random_state=42)
X_train, X_val, y_train, y_val = train_test_split(X_train, y_train, test_size=0.2, random_state=42) # 0.2 x 0.75 = 0.15
```

<!-- grade_id: cell-8da00aed9ea7c071, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == -0.08024
assert round(X_val[0][8], 5) == 0.09924
assert round(X_test[0][8], 5) == 0.03243
```

**6** Υλοποιήστε συνάρτηση, η οποία δοθέντος ενός δενδρικού μοντέλου παλινδρόμησης, διερευνά όλες τις ακέραιες τιμές της μεταβλητής min_samples_leaf από ένα κάτω όριο μέχρι ένα άνω όριο και επιστρέφει την βέλτιστη με βάση το mean squared error στο παραπάνω σύνολο επικύρωσης έπειτα από εκπαίδευση στο παραπάνω σύνολο εκπαίδευσης. Επιπλέον επιστρέφει και το mean squared error (στρογγυλοποιημένο σε ακέραιο) που πετυχαίνει στο σύνολο ελέγχου το μοντέλο με την βέλτιστη επίδοση στο σύνολο επικύρωσης. (2 μονάδες)

<!-- grade_id: cell-45e2ac4f9c5115f9, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import mean_squared_error

def best_min_samples_leaf(model, min, max):
    # Initialize a variable to hold the best min samples leaf, as well as the best error found both in validation and test sets
    best_msl = 0
    best_val_mse = float('inf')
    best_test_mse = float('inf')
    
    # Iterate over all possible min sample leaf values
    for msl in range(min, max + 1):
        # Get a copy of the model, with the parameter min sample leaf set
        this_model = model.set_params(min_samples_leaf=msl)
        # Fit the model to the training data, and make predictions on the validation set
        this_model.fit(X_train, y_train)
        y_pred = this_model.predict(X_val)
        # Calculate MSE metrics on test and validation set
        val_mse = mean_squared_error(y_val, y_pred)
        test_y_pred = this_model.predict(X_test)
        test_mse = mean_squared_error(y_test, test_y_pred)
        
        # Keep the metrics for the best validation MSE
        if val_mse < best_val_mse:
            best_val_mse = val_mse
            best_msl = msl
            best_test_mse = test_mse
    
    # Return the best min sample leaf value, along with the MSE on the test set for that value
    return (best_msl, round(best_test_mse))
```

<!-- grade_id: cell-0d6be21e54be9665, points: 2, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_depth=3, random_state=42)
assert best_min_samples_leaf(model, 1, 8) == (7, 3358)
assert best_min_samples_leaf(model, 5, 15) == (11, 3271)
```

**7** Υλοποιήστε συνάρτηση η οποία δοθέντος ενός μοντέλου παλινδρόμησης, ενός συνόλου δεδομένων και μίας μετρικής, το αξιολογεί με τη μέθοδο της σταυρωτής επικύρωσης "άφησε ένα εκτός". (1 μονάδα)

<!-- grade_id: cell-880c64392f85bccb, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import LeaveOneOut
import numpy as np

def leave_one_out(model, X, y, metric):
    # Initialize leave one out splitter
    loo = LeaveOneOut()
    scores = []
    
    # Iterate over all folds
    for train_idx, val_idx in loo.split(X):
        # Use separate model each time to ensure tests are independent
        this_model = model
        X_train, X_val = X[train_idx], X[val_idx]
        y_train, y_val = y[train_idx], y[val_idx]
        
        # Fit the model to the training data, and make predictions on the test set
        this_model.fit(X_train, y_train)
        y_pred = this_model.predict(X_val)
        
        # Calculate score via the provided metric function
        score = metric(y_val, y_pred)
        
        # Add this score to the list
        scores.append(score)
    
    # Return mean of all scores, evaluating the regressor
    return np.mean(scores)
```

<!-- grade_id: cell-4afd7b6de1cd1f7f, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_leaf_nodes=6, min_samples_leaf=2, random_state=42)
assert round(leave_one_out(model, diabetes.data, diabetes.target, mean_squared_error), 0) == 3354
```